rule_files:
  - streams.yml
evaluation_interval: 1m
tests:
- interval: 1m
  input_series:
  - series: 'ws_gap_ms_bucket{le="1500",exchange="binance",channel="aggTrade"}'
    values:  [1 1 1 1 1 1 1 1 1 1 1]
  - series: 'ws_gap_ms_bucket{le="+Inf",exchange="binance",channel="aggTrade"}'
    values:  [2 2 2 2 2 2 2 2 2 2 2]
  alert_rule_test:
  - eval_time: 10m
    alertname: StreamsLagHigh
    exp_alerts: []   # under 1500ms, should NOT fire
- interval: 1m
  input_series:
  - series: 'ws_gap_ms_bucket{le="1500",exchange="binance",channel="aggTrade"}'
    values:  [0 0 0 0 0 0 0 0 0 0 0]
  - series: 'ws_gap_ms_bucket{le="+Inf",exchange="binance",channel="aggTrade"}'
    values:  [5 5 5 5 5 5 5 5 5 5 5]
  alert_rule_test:
  - eval_time: 10m
    alertname: StreamsLagHigh
    exp_alerts:
    - exp_labels:
        exchange: binance
        channel: aggTrade
